# Project Euler: Problem 52

https://projecteuler.net/problem=52

It can be seen that the number, $125874$, and its double, $251748$, contain exactly the same digits, but in a different order.

Find the smallest positive integer, $x$, such that $2x$, $3x$, $4x$, $5x$, and $6x$, contain the same digits.

First, write a function `same_digits(x,y)` that returns `True` if two integers `x` and `y` have the exact same set of digits and multiplicities and `False` if they have different digits.

```python

```

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def same_digits(x, y):
    """Do the integers x and y have the same digits, regardless of order."""
    '''if len(x)==len(y):
        a=True
    for i in x:
        if i in y:
            b=True
    if a==True and b==True:
        return True'''
    X=[]
    Y=[]
    for i in x:
        X.append(int(i))
    for i in y:
        Y.append(int(i))
    #put X and Y in order
    if X==Y:
        return True
    '''split x and y into a list X and Y
    use function to put the list in numerical order
    if X==Y return True'''
```

```python
same_digits('537653','82462846')
```

Output:
```
([5, 3, 7, 6, 5, 3], [8, 2, 4, 6, 2, 8, 4, 6])
```

<!-- grade_id: projecteuler52a, points: 4, grade: true, solution: false -->
```python
assert same_digits('132', '321')
assert not same_digits('123', '3')
assert not same_digits('456', '0987654321')
```

Output:
```
<ipython-input-5-4d2369f19b4d> in <module>()
      1 assert same_digits('132', '321')
----> 2 assert not same_digits('123', '3')
      3 assert not same_digits('456', '0987654321')

<ipython-input-4-575ba1876147> in same_digits(x, y)
      6         if i in y:
      7             b=True
----> 8     if a==True and b==True:
      9         return True

UnboundLocalError: local variable 'a' referenced before assignment
```

Now use the `same_digits` function to solve this Euler problem. As you work on this problem, be careful to debug and test your code on small integers before trying it on the full search.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: projecteuler52b, points: 6, grade: true, solution: false -->
```python
assert True # leave this cell to grade the solution
```